# Review Module

**Instructions**

For this review module, the exercises will require you to uncomment lines of code. In order to complete this part, we recommend you follow these instructions:

1. Do *not* change anything other than removing the '#' before the answer you choose. If you change anything else, the autograder will fail and you will not receive any points.
2. To check that your answers are correct, click on `Kernel > Restart & Run All`. If any of the cells in your notebook produces an error, that means you need to adjust your answers.

### Exercise 1

In the Laundry/Closet example, which of the following would represent a data warehouse?

* A. The closet itself
* B. A drawer in the closet
* C. The laundry basket
* D. The outfit that is assembled

<!-- grade_id: cell-1e8ffcaefc5829c6, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
# raise NotImplementedError() # Remove this line when you enter your solution

ex_1_answer = "A"
#ex_1_answer = "B"
#ex_1_answer = "C"
#ex_1_answer = "D"
```

### Exercise 2

Which type of relational database dimension modeling mimics the entity relationship way of data modeling, where all entities are connected to one table of facts?

* A. OLAP Cubes
* B. Star Schema
* C. Non-relational Schema
* D. Fact Table

<!-- grade_id: cell-0da2b6bc083c5bfe, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
# raise NotImplementedError() # Remove this line when you enter your solution

#ex_2_answer = "A"
ex_2_answer = "B"
#ex_2_answer = "C"
#ex_2_answer = "D"
```

### Exercise 3

In Amazon Redshift, what is a data warehouse called?

* A. node
* B. bucket
* C. instance
* D. cluster

<!-- grade_id: cell-3d1ce496bafa6e17, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
# raise NotImplementedError() # Remove this line when you enter your solution

#ex_3_answer = "A"
#ex_3_answer = "B"
#ex_3_answer = "C"
ex_3_answer = "D"
```

### Exercise 4

In an organization, what would you use to set specific permissions for different users, such as giving someone AmazonRedshiftFullAccess?

* A. Root
* B. S3
* C. IAM
* D. ARN

<!-- grade_id: cell-0a9f8491399dad70, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
# raise NotImplementedError() # Remove this line when you enter your solution

#ex_4_answer = "A"
#ex_4_answer = "B"
ex_4_answer = "C"
#ex_4_answer = "D"
```

### Exercise 5

Which of these is cloud storage where we could create a bucket as a temporary stop for our data before we load and transform it?

* A. Root
* B. S3
* C. IAM
* D. ARN

<!-- grade_id: cell-5b8bce01ff596ccf, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
# raise NotImplementedError() # Remove this line when you enter your solution

#ex_5_answer = "A"
ex_5_answer = "B"
#ex_5_answer = "C"
#ex_5_answer = "D"
```

## Testing Cells

Run the below cells to check your answers. Make sure you run your solution cells first before running the cells below, otherwise you will get a `NameError` when checking your answers.

<!-- grade_id: cell-f2406da504968cd9, points: 1, grade: true, solution: false -->
```python
# Ex. 1
assert "ex_1_answer" in globals(), "Ex. 1 - You haven't defined the ex_1_answer variable!"
print("Preliminary sanity check for Ex. 1 passed.")
assert ord(ex_1_answer) == 65, "Ex. 1 - Think again! Check the main case for a refresher!"
print("Exercise 1 seems correct!")
```

Output:
```
Preliminary sanity check for Ex. 1 passed.
Exercise 1 seems correct!
```

<!-- grade_id: cell-a584e0fa552c4325, points: 1, grade: true, solution: false -->
```python
# Ex. 2
assert "ex_2_answer" in globals(), "Ex. 2 - You haven't defined the ex_1_answer variable!"
print("Preliminary sanity check for Ex. 2 passed.")
assert ord(ex_2_answer) == 66, "Ex. 2 - Think again! Check the main case for a refresher!"
print("Exercise 2 seems correct!")
```

Output:
```
Preliminary sanity check for Ex. 2 passed.
Exercise 2 seems correct!
```

<!-- grade_id: cell-05097b8377476302, points: 1, grade: true, solution: false -->
```python
# Ex. 3
assert "ex_3_answer" in globals(), "Ex. 3 - You haven't defined the ex_3_answer variable!"
print("Preliminary sanity check for Ex. 3 passed.")
assert ord(ex_3_answer) == 68, "Ex. 3 - Think again! Check the main case for a refresher!"
print("Exercise 3 seems correct!")
```

Output:
```
Preliminary sanity check for Ex. 3 passed.
Exercise 3 seems correct!
```

<!-- grade_id: cell-a5f826387be825a0, points: 1, grade: true, solution: false -->
```python
# Ex. 4
assert "ex_4_answer" in globals(), "Ex. 4 - You haven't defined the ex_4_answer variable!"
print("Preliminary sanity check for Ex. 4 passed.")
assert ord(ex_4_answer) == 67, "Ex. 4 - Think again! Check the main case for a refresher!"
print("Exercise 4 seems correct!")
```

Output:
```
Preliminary sanity check for Ex. 4 passed.
Exercise 4 seems correct!
```

<!-- grade_id: cell-75f03680e0531c84, points: 1, grade: true, solution: false -->
```python
# Ex. 5
assert "ex_5_answer" in globals(), "Ex. 5 - You haven't defined the ex_5_answer variable!"
print("Preliminary sanity check for Ex. 5 passed.")
assert ord(ex_5_answer) == 66, "Ex. 5 - Think again! Check the main case for a refresher!"
print("Exercise 5 seems correct!")
```

Output:
```
Preliminary sanity check for Ex. 5 passed.
Exercise 5 seems correct!
```